Food-washing monkeys recognize the law of diminishing returns

Author response:

The following is the authors’ response to the original reviews.

We thank the reviewers for their constructive criticism. It is rare and gratifying to receive such thoughtful feedback, and the result is a much stronger paper. We made significant changes to our statistical analyses and figures to better differentiate the effects of sex and dominance rank on food-cleaning behaviors. These revisions uphold our original conclusion––that rank-related variation overwhelms any sex difference in cleaning behavior. We hope that these edits, together with the rest of our responses, provide a convincing demonstration of the tradeoffs of eliminating quartz from food surfaces.

Reviewer #1 (Public Review):

Summary

We have no objections to Reviewer 1’s summary of our manuscript.

Strengths

Reviewer 1 is extremely gracious, and we are grateful for the kind words.

Weaknesses

Reviewer 1 identified several weaknesses, enumerating three types: (1) statistics, (2) insufficient links to foraging theory, and (3) interpretation and validity of the model. The present response is organized around these same categories.

(1) Statistics

We put all of our data and code into the Zenodo repository prior to submission. This content should have been accessible to Reviewer 1 from the outset. But in any event, we are very sorry for the mixup. To ensure access to our data and code during the present stage of review, we included the URL in the main mainscript and here: https://doi.org/10.5281/zenodo.14002737

(a) AIC and outcome distributions

Reviewer 1 criticized our use of AIC for determining model selection. We agree and this aspect of our manuscript is now removed. In lieu of AIC, we produced two data sets consisting of whole number counts (seconds) with means <5. The data were right-skewed due to high concentrations of biologically-meaningful zeros (i.e., bouts of food handling without any cleaning effort). Following the recommendations of Bolker et al. (2008) and others (Brooks et al. 2017, 2019), we chose an outcome distribution (zero-inflated Poisson, see response below) that best matched this data distribution. In addition, we evaluated the post-hoc performance of each of our models using the standardized residual diagnostic tools for hierarchical regression models available in the DHARMa package (Hartig, 2022). To further evaluate our choice of outcome distribution, we generated QQ-plots and residual vs. predicted plots for each model and included them in our revision as Figures S3-S5.

(b) zeros

Reviewer 1 expressed concern over our treatment of biologically-meaningful zeros, and recommended use of a zero-inflated GLMM with either a Poisson or negative binomial outcome distribution. We agree that such models are best for our two data sets. Accordingly, we fit a series of zero-inflated generalized linear mixed models (ZIGLMM) using the glmmTMB package in R, each with a logit-link function, a single zero-inflation parameter applying to all observations, and a Poisson error distribution. For the food-brushing model, we fit a zero-inflated Poisson (ZIP), which produced favorable standardized residual diagnostic plots with no major patterns of deviation (Figure S3) and minor, but non-significant underdispersion (DHARMa dispersion statistic = 0.99, p = 0.80). For our two food-washing models, we used zero-inflated models with Conway-Maxwell Poisson (ZICMP) distributions, an error distribution chosen for its ability to handle data that are more underdispersed (DHARMa dispersion statistic = 8.2E-09, p = 0.74) than the standard zero-inflated Poisson (Brooks et al. 2019). Using this error distribution improved residual diagnostic plots over a standard ZIP model and we view any deviations in the standardized residuals as minor and attributable to the smaller sample size of our food-washing data set (see Figures S4 and S5) (Hartig, 2022). We reported the summarized fixed effects tests for each GLMM in Tables S1-S3 as Analysis of Deviance Tables (Type II Wald chi square tests, one-sided) along with 𝜒2 values, degrees of freedom, and p-values (one-sided tests). Full model summaries with standard errors and confidence intervals are also included in Tables S4-S6. For all statistical analyses, we set 𝛼 = 0.05.

(2) Absence of Links to Foraging Theory

This critique has three components. The first revisits the absence of code for the optimal cleaning time model. This omission was an unfortunate error at the moment of submission, but our code is available now as a Mathematica notebook in Zenodo (https://doi.org/10.5281/zenodo.14002737). The second pivots around our scholarship, admonishing us for failing to acknowledge the marginal value theorem of Charnov (1976). It is a fair point and we have corrected the oversight with a citation to this classic paper. The third criticism is also rooted in scholarship, with Reviewer 1 asking for greater connection to the existing literature on optimal foraging theory, a point echoed in the summary assessment of the editors at eLife. This comment and the weight given to it by eLife’s editors put us in a difficult spot, as our paper is focused on the optimization of delayed gratification, not food acquisition per se. So, we are in the awkward position of gently resisting this recommendation while simultaneously agreeing with Reviewer 1 that we need to better situate our findings in the landscape of existing literature. To thread this needle, we produced Box 2 with a photograph and 410 words. This display box puts our findings into direct conversation with recent research focused on the sunk cost fallacy.

(3) Interpretation and validity of model relative to data

This critique is focused on the simulated brushing and washing results reported in Figure S1, along with its captioning, which was inadequate. We edited the caption to identify the author (JER) who simulated the brushing and washing behaviors of the monkeys. In addition, we clarified the number of brushing replicates (3) and washing replicates (3) for each of three treatments, for a total of 18 simulations.

We followed Reviewer 1’s suggestion, incorporating the experimental uncertainty of grit removal into our optimal cleaning time model. We drew % grit removed values the % grit removed is used to estimate the cleaning inefficiency≥ 100%parameter 𝑐 for from a distribution, discounting the rare event when values were drawn. As brushing and washing, the included uncertainty now allows us to evaluate these parameters as distributions; and, in turn, obtain a distribution for our predicted brushing and washing optimal cleaning times. As we now describe in the main text, the optimal cleaning time for brushing and washing are 𝑡* = 0. 98 ± 0. 19 s and * = 2. 40 ± 0. 74 s, respectively. We are grateful for Reviewer 1’s suggestion, for it added𝑡 valuable context to our model predictions. Notably, the inclusion of experimental uncertainty did not change the qualitative nature of our results, or the interpretations of our model predictions compared to observed cleaning behaviors.

We choose to exclude variability in handling time h to generate predicted cleaning time optima, at least in the main text. Our reasoning stems from the observation that handling time variability is long-tailed, with the longer handling times associated with behaviors that we do not account for in our analysis. For example, individuals carrying multiple cucumber slices to the ocean were apt to drop them, struggling at times to re-grasp so many at once. Such moments increased handling times substantially. Still, we acted on Reviewer 1’s suggestion, accounting for the tandem effects of handling time variability and uncertainty in % grit removed (see Figure S6). Drawing handling time estimates from a log-normal distribution fitted to the handling time data, we found that these dual sources of uncertainty did not qualitatively change our results. They added further uncertainty to the predicted washing time, but the mean remains roughly equivalent. (We note that brushing is assumed to have a constant handling time––composed of only assessment time and no travel––such that the results for brushing do not change.) Both analyses are included in the Mathematica notebook at (https://doi.org/10.5281/zenodo.14002737).

Reviewer #2 (Public Review):

Summary

We have no objections to Reviewer 2’s summary of our manuscript.

Strengths

Reviewer 2 is extremely gracious, and we are grateful for the kind words.

Weaknesses

Reviewer 2 noted that our manuscript failed to provide “sufficient background on [our study] population of animals and their prior demonstrations of food-cleaning behavior or other object-handling behaviors (e.g., stone handling).” To address this comment, we edited the introduction (lines 56-58) to alert readers to the onset of regular food-cleaning behaviors sometime after December 26, 2004. In addition, we edited our methods text (lines 155-160) to highlight the onset and limited scope of prior research with this study population:

“The animals are well habituated to human observers due to regular tourism and sustained study since 2013 (Tan et al., 2018). Most of this research has revolved around stone tool-mediated foraging on mollusks, the only activity known to elicit stone handling (Malaivijitnond et al., 2007; Gumert and Malaivijitnond, 2012, 2013; Tan et al., 2015), although infants and juveniles will sometimes use stones during object play (Tan, 2017). There has been no prior examination of food-cleaning behaviors.”

Reviewer #3 (Public Review):

Reviewer 3 identified three weaknesses, which we address in three paragraphs.

Reviewer 3 questioned our methods for determining rank-dependent differences in cleaning behavior, arguing that our conclusions were unsupported. It is a fair point, and it compelled us to combine males and females into a single standardized ordinal rank of 24 individuals. This unified ranking is now reflected in the x-axes of Figure 2 and Figure S2. Plotting the data this way––see Figure S2––underscores Reviewer 3’s concern that sex and dominance rank are confounding variables. To address this problem, our GLMM included rank and sex as predictor variables, which controls for the effect of sex when assessing the relationship between rank and cleaning time across the three treatments. Reported in Tables S1-S3, these findings show that the effect of sex on either brushing or washing time was not significant. This result bolsters our original contention that rank-related variation in cleaning time overwhelms any sex differences.

Relatedly, Reviewer 3 questioned our conclusions on the effects of rank because our study was focused on a single social group. In other words, it is plausible that our results were heavily influenced by the idiosyncrasies of select individuals, not dominance rank per se. It is a fair point, and it compelled us to include individual ID as a random effect in each of our GLMMs. Including individual ID as a random intercept allowed us to control for inter-individual variation in cleaning duration while assessing the effects of rank. An analysis based on additional social groups or longitudinal data are certainly desirable, but also well beyond the scope of a Short Report for eLife.

Finally, Reviewer 3 objected to fragments of sentences in our abstract, introduction, and discussion, combining them into a criticism of claims that we did not and do not make. It probably wasn’t intentional, but it puts us in the awkward position of deconstructing a strawman:

● Review 3 begins, “there is no evidence presented on the actual fitness-related costs of tooth wear or the benefits of slightly faster food consumption”. This statement is true while insinuating that collecting such evidence was our intent. To be clear, our experiment was never designed to measure tooth wear or reproductive fitness, nor do we make any claims of having done so.

● Reviewer 3 adds, “Support for these arguments is provided based on other papers, some of which come from highly resource-limited populations (and different species). But this is a population that is supplemented by tourists with melons, cucumbers, and pineapples!” We were puzzled over these sentences. The first fails to mention that the citations exist in our discussion. Citing relevant work in a discussion is a basic convention of scientific writing. But it seems the underlying intent of these words is to denigrate the value of our study population because two dozen tourists visit Koram Island once a day. Exclamations to the contrary, the amount of tourist-provisioned food in the diet of any one monkey is negligible.

● Last, Reviewer 3 commented on matters of style, objecting to “overly strong claims.” We puzzled over this criticism because the claims in question are broader points of introduction or discussion, not results. The root problem appears to be the final sentence of our abstract:

“Dominant monkeys abstained from washing, balancing the long-term benefits of mitigating tooth wear against immediate energetic requirements, an essential predictor of reproductive fitness.”

This sentence has three clauses. The first is a statement of results, whereas the second and third are meant to mirror our discussion on the importance of our findings. We combined the concepts into a single concluding sentence for the sake of concision, but we can appreciate how a reader could feel deceived, expecting to see data on tooth wear and fitness. So, our impression is that we are dealing with a simple misunderstanding of our own making, and that this single sentence explains Reviewer 3’s criticism and tone––it cast a long shadow over the substance of our paper. To resolve this problem, we edited the sentence:

“Dominant monkeys abstained from washing, a choice consistent with the impulses of dominant monkeys elsewhere: to prioritize rapid food intake and greater reproductive fitness over the long-term benefits of prolonging tooth function.”

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Few animal species have the cognitive faculties or prehensile abilities needed to eliminate costly tooth-damaging grit from food surfaces. Some populations of monkeys wash sand from foods when standing water is readily accessible, but this propensity varies within groups for reasons unknown. Spontaneous food-washing emerged recently in a group of long-tailed macaques (Macaca fascicularis) inhabiting Koram Island, Thailand, and it motivated us to explore the factors that drive individual variability. We measured the mineral and physical properties of contaminant sands and conducted a field experiment, eliciting 1,282 food-handling bouts by 42 monkeys. Our results verify two long-standing presumptions, that monkeys have a strong aversion to sand and that removing it is intentional. Reinforcing this result, we found that monkeys clean foods beyond the point of diminishing returns, a suboptimal behavior that varies with rank. Dominant monkeys abstain from washing, balancing the long-term benefits of mitigating tooth wear against immediate energetic requirements, an essential predictor of reproductive fitness.

Revised version: Few animals have the cognitive faculties or prehensile abilities needed to eliminate tooth-damaging grit from food surfaces. Some populations of monkeys wash sand from foods when standing water is readily accessible, but this propensity varies within groups for reasons unknown. Spontaneous food-washing emerged recently in a group of long-tailed macaques (Macaca fascicularis) inhabiting Koram Island, Thailand, and it motivated us to explore the factors that drive individual variability. We measured the mineral and physical properties of contaminant sands and conducted a field experiment, eliciting 1,282 food-handling bouts by 42 monkeys. Our results verify two long-standing presumptions, that monkeys have a strong aversion to sand and that removing it is intentional. Reinforcing this result, we found that monkeys clean foods beyond the point of diminishing returns, a suboptimal behavior that varied with social rank. Dominant monkeys abstained from washing, a choice consistent with the impulses of dominant monkeys elsewhere: to prioritize rapid food intake and greater reproductive fitness over the long-term benefits of prolonging tooth function.

Introduction

Earlier version: Koshima Island, Japan is a storied fieldsite in the annals of primatology. Observations of wild macaques (Macaca fuscata) began there in 1948, but four years of sustained effort failed to habituate the monkeys. In August 1952, Itani and Tokuda (1954) resorted to scattering wheat grains and sweet potatoes on oft-used paths, gradually shifting the provisions to a sandy beach at Otomari Bay, a strategy that afforded a clear view of the entire group of 22 animals. Individual identifications followed quickly, laying the foundation for decades of influential research. In September 1953, a young female named Imo gathered a sweet potato from the beach and washed it in a freshwater stream, a behavioral innovation that spread horizontally to peers and then vertically to older kin (Kawai, 1965). By 1958, sweet potato washing had become a group-wide trait with a key modification: the monkeys began using sea water instead of standing freshwater, a preference that continues today seven generations later (Hirata et al., 2001).

Revised version: Koshima Island, Japan is a storied fieldsite in the annals of primatology. Observations of wild macaques (Macaca fuscata) began in 1948, but four years of sustained effort failed to habituate the monkeys. In August 1952, Itani and Tokuda (1954) resorted to scattering wheat grains and sweet potatoes along oft-used paths, gradually shifting the provisions to a sandy beach at Otomari Bay, a strategy that afforded a clear view of the entire group of 22 animals. Individual identifications followed quickly, laying the foundation for decades of influential research. In September 1953, a young female named Imo gathered a sweet potato from the beach and rinsed it in a freshwater stream, a behavioral innovation that spread horizontally to peers and then vertically to older kin (Kawai, 1965). By 1958, sweet potato washing had become a group-wide trait with a key modification: the monkeys began using sea water instead of standing freshwater, a preference that continues today seven generations later (Hirata et al., 2001).

Earlier version: These events have since passed into canon as an example of socially transmitted behavior, or culture, among nonhuman primates (McGrew, 1998; Matsuzawa and McGrew, 2008; Matsuzawa, 2015). But our preoccupation with culture has overshadowed fundamental questions of motivation (Sarabian and MacIntosh, 2015; Fiore et al., 2020). Two tacit assumptions—that sand produces an objectionable sensation on teeth, and that it is prudent to minimize tooth damage—are sufficiently intuitive that formal tests are wanting (Fannin et al., 2021). In consequence, the mineral and physical properties of contaminant sands are unknown, let alone the efficiency of different cleaning behaviors (Schofield et al., 2018). Another enigma concerns the preference of some monkeys to brush food on themselves (Kawai et al., 1992), a rapid but seemingly inferior means of sand removal. Food-brushing individuals have been characterised as inept (Kawai, 1965) or subordinate (Watanabe, 1994) in part because the quartz in sand can cause severe tooth damage (Lucas et al., 2013; Towle et al., 2022). Yet, carrying food to the ocean is expected to incur energetic costs as well as opportunity costs, factors that impelled usto explore the trade-offs of mitigating sand-mediated tooth wear.

Revised version: These events have since passed into canon as an example of socially transmitted behavior, or culture, among nonhuman primates (McGrew, 1998; Matsuzawa and McGrew, 2008; Matsuzawa, 2015). It is a venerable legacy, but it overshadows a fundamental question: why do monkeys wash their food? (Sarabian and MacIntosh, 2015; Fiore et al., 2020). Two tacit assumptions—that sand produces an objectionable sensation on teeth, and that it is prudent to minimize tooth damage—are sufficiently intuitive that formal tests are wanting. In consequence, the mineral and physical properties of contaminant sands are unknown, let alone the efficiency of different cleaning behaviors (Schofield et al., 2018). Another enigma concerns the preference of some monkeys to brush food with their hands (Kawai et al., 1992), a rapid but seemingly inferior means of sand removal. Foodbrushing individuals have been characterized as inept (Kawai, 1965) or subordinate (Watanabe, 1994) in part because the quartz in sand can cause severe tooth damage (Lucas et al., 2013; Towle et al., 2022). Yet, carrying food to the ocean is expected to incur energetic costs and opportunity costs, factors that impelled us to explore the trade-offs of mitigating sand-mediated tooth wear.

Koram Island, Thailand

Earlier version: The long-tailed macaques (M. fascicularis) of Koram Island, Thailand use stone tools to harvest shellfish, a phenomenon that came to light during biodiversity damage assessments following the Sumatra-Andaman earthquake and tsunami of December 26, 2004 (Malaivijitnond et al., 2007; Gumert and Malaivijitnond, 2012; Tan et al., 2015). The monkeys are now a magnet for tourists, who provide them with market-sourced fruits (cucumbers, melon, pineapple), jettisoning them onto the beach. These events produce food surfaces with considerable concentrations of sand ( x¯ = 3.7 ± 1.3 mg mm-2) and elicit food-washing and food-brushing behaviors among the monkeys. To understand the factors driving these reactions, we examined the mineral properties of foodadhering sands (n = 758 particles), finding that 78% of our sample was composed of crystalline quartz (Figure 1A).

Revised version: The long-tailed macaques (M. fascicularis) of Koram Island, Thailand use stone tools to harvest shellfish, a phenomenon that came to light during the surveys of biodiversity damage that followed the Sumatra-Andaman earthquake and tsunami of December 26, 2004 (Malaivijitnond et al., 2007; Gumert and Malaivijitnond, 2012; Tan et al., 2015). The monkeys then became a magnet for tourism. Some visitors began supplying the monkeys with market-sourced fruits (cucumbers, melon, pineapple), jettisoning them onto the beach. These events produce food surfaces with considerable concentrations of sand ( x¯=3.7±1.3 mg mm−2 ), which, in turn, elicit food-washing and food-brushing behaviors among the monkeys. To understand the factors driving these reactions, we examined the mineral properties of food-adhering sands (n = 758 particles), finding that 78% of our sample was composed of crystalline quartz (Figure 1A).

Variation in the mineral and physical properties of quartz particles on food surfaces.

(A) Particle sizes followed a bimodal distribution, with most particles featuring metal inclusions. Nearly half the sample is < 25 μm, the “grittiness threshold” of the human oral cavity (Imai et al., 1995). (B) Circularity is a dimensionless shape factor (range: 0 to 1) based on two-dimensional microscopy and estimates for the projected area and perimeter of a given particle. Some examples are illustrated; overall, it is a convenient but imperfect proxy for sphericity, or deviations from spherical (Grace and Ebneyamini, 2021). Here, circularity varied as a function of mean Feret diameter, suggesting that larger particles hold greater potential for damaging attack angles during particle-enamel contact.

Earlier version: Harder than enamel, quartz can exact a heavy toll on teeth, but the probability and degree of enamel loss is governed partly by the size and shape of individual particles, factors that determine the ‘attack angle’ during particle-enamel contact (Lucas et al., 2013). We calculated the circularity of particles as a convenient proxy for sphericity (Grace and Ebneyamini, 2021), finding that it decreased as a function of particle size (Figure 1B). This result suggests that larger particles are more angular, posing a greater risk to enamel. However, we also calculated a median Feret diameter of 25.8 μm (range: 8.7 to 644 μm), meaning that nearly half the sample existed below the human threshold (25 μm) of oral detection (Imai et al., 1995). To put 25 pm into perspective, it is one-twelfth the diameter of the period ending this sentence.

Revised version: Harder than enamel, quartz can exact a heavy toll on teeth, but the probability and degree of enamel loss is governed partly by the size and shape of individual particles, factors that determine the ‘attack angle’ during particle-enamel contact (Lucas et al., 2013). We calculated the circularity of particles as a convenient proxy for sphericity (Grace and Ebneyamini, 2021), finding that it decreased as a function of particle size (Figure 1B). This result suggests that larger particles are more angular, posing a greater risk to enamel. However, we also calculated a median Feret diameter of 25.8 μm (range: 8.7 to 644 μm), meaning that nearly half the sample existed below the human threshold (25 μm) of oral detection (Imai et al., 1995). To put 25 μm into perspective, it is one-twelfth the diameter of the period ending this sentence.

Mitigating tooth wear

Chewing undetected quartz is expected to cause severe tooth wear, but this cost can be mitigated behaviorally if food-cleaning is proficient. To test this contention, we simulated the brushing and washing actions of monkeys with cucumber slices exposed to three concentrations of sand: low ( x¯=0.2±0.2 mg mm−2 ), intermediate ( x¯=0.9±0.1 mg mm−2 ), and high ( x¯=1.8±0.9 mg mm−2 ). We found that brushing was less efficient than washing across treatments, eliminating 76 ± 7% vs. 93 ± 4% of sand particles, respectively (Figure S1). It is a modest difference, perhaps, but it is freighted with fitness consequences when extrapolated over years of life (Fannin et al., 2022). It follows that monkeys should compulsively wash sand from food whenever the opportunity avails itself, a prediction that motivated a field experiment.

Results and discussion

Earlier version: Our experiment (Figure 2A) was designed to test two concepts at once. The first pivots around intentionality, a thorny problem that emerged from studies of raccoons (Procyon lotor). Celebrated food-handlers, the submersion of food objects in water is better termed ‘dousing’ for greater haptic sensation, not washing with the intention of removing surface contaminants (Lyall-Watson, 1963). If the intent of monkeys is to eliminate sand, then the time devoted to brushing or washing food should vary as a positive function of sandiness. The other concept draws on observations from Koshima Island, which alluded to rank effects on individual cleaning behaviors, a pattern that is difficult to detect without controlling access to food or distance to the ocean.

Revised version: Our experiment was designed to test two concepts at once (Figure 2; Video 1). The first pivots around intentionality, a thorny problem that emerged from studies of raccoons (Procyon lotor). Celebrated food-handlers, the submersion of food objects in water is better termed ‘dousing’ for greater haptic sensation, not washing with the intention of removing surface contaminants (Box 1). If the intent of monkeys is to eliminate sand, then the time devoted to brushing or washing food should vary as a positive function of sandiness. The other concept draws on observations from Koshima Island, which alluded to rank effects on individual cleaning behaviors, a pattern that is difficult to detect without controlling access to food or distance to the ocean.

Experimental design and results.

Earlier version: (A) To elicit food-cleaning behaviors, we put sliced cucumbers in trays representing three treatments—food surfaces with low ( x¯ = 0.2 ± 0.2 mg mm-2), intermediate ( x¯ = 0.9 ± 0.1 mg mm-2), and high ( x¯ =1.8 ± 0.9 mg mm-2) concentrations of sand—positioned 1.5 m apart and 15 m from the ocean. (B) Monkeys brushed sandier treatments for longer durations, χ2 (2, n = 575 food-handling bouts) = 185.7, p < 0.0001) with no difference across dominance ranks (Table S3). (C) Monkeys washed sandier treatments for longer durations, χ2 (2, n = 362 food-handling bouts) = 50.5, p < 0.0001), and we detected an interaction effect with dominance rank, χ2 = 6.4, p = 0.04 (Table S4). Raw data for panels B and C are illustrated in Figure S2. (D) Energy intake rates also varied as function dominance (ANOVA; F2,104 = 4.3; p = 0.02). Symbols represent mean values and whiskers ± 1 s.e. Photos by Amanda Tan.

Revised version: (A) To elicit food-cleaning behaviors, we put sliced cucumbers in trays representing three treatments—food surfaces with low ( x¯=0.2±0.2 mg mm−2 ), intermediate ( x¯=0.9±0.1 mg mm−2 ), and high ( x¯=1.8±0.9 mg mm−2 ) concentrations of sand—positioned 1.5 m apart and 15 m from the ocean. (B) Monkeys brushed the sandier treatments for longer durations [χ2 (2, n = 575 food-handling bouts) = 194.7, p < 0.0001] with no effect of dominance rank or sex (Table S1; Figure S2). (C) Monkeys washed the sandier treatments for longer durations [χ2 (2, n = 362 food-handling bouts) = 69.7, p < 0.0001], and we found an interaction effect with rank independent of sex [χ2 = 19.3, p < 0.0001; Table S2; Figure S2]. (D) Energy intake rates also varied as function dominance rank [ANOVA, (LN-transformed); F2,104 = 10.0; p < 0.0001]. Symbols represent mean values and whiskers ± 1 s.e. Photos by Amanda Tan.

Added in the revised version: Box 1. Manual prehension and the nomenclature of raccoons

Added in the revised version: Raccoons (Procyon lotor) tend to live in wooded habitats near waterways, where they can be seen dousing foods before ingestion (photograph by Markus Zindl, reproduced with permission).

Added in the revised version: Animal names tend to reflect salient features of their appearance or behavior, including the sounds they make. This basic principle of ethnotaxonomy harmonizes the classification of animals with their traits, a pattern that extends to other languages when names are borrowed or translated. But this process can have profound consequences, shaping our understanding of animals and their behaviors. For example, manipulative actions are baked into the word raccoon, a corruption of the Powhatan words arakun or arakunem, meaning, approximately, ‘it scratches with its hands.’ Far more common across North America is the Anishinaabe word esiban (‘it picks up things’), with cognates in the languages of at least a dozen cultural groups, including Cree, Potawatomi, Abenaki, and Delaware peoples (Holmgren, 1990; Justice, 2021). The Tsimshian word que-o-koo (‘washes with hands’) is notably different for imputing intent, perhaps because coastal peoples are sensitized to sand on their foods.

Added in the revised version: Linnaeus put raccoons in the family Ursidae—bears—in the second edition of Systema Naturae (1740), classifying them as Ursus cauda elongata (‘bear with long tail’). By the 10th edition (1758), he had reclassified them as Ursus lotor (‘washer bear’) due to accounts of captive raccoons persistently dunking food in water, along with his own observations of Sjupp, a pet raccoon gifted to him by the Swedish crown prince Adolf Fredrik (Nicholls, 2007). Twenty years later, in 1780, the German naturalist Gottlieb Conrad Christian Storr elevated raccoons into their own genus, Procyon, meaning’before the dog’or’early dog’, a nod to their dog-like appearance;but he retained the species nomen lotor. Today, the raccoon is known as tvättbjörn in Linnaeus’s Swedish and waschbär in Storr’s German, both meaning ‘washer bear.’

Added in the revised version: Meanwhile, the French naturalist Georges-Louis Leclerc, Comte de Buffon conducted his own detailed observations of pet raccoons held in the Muséum National d’Histoire Naturelle, Paris. He saw affinities with rodents, a legacy that echoes today in the names raton laveur and ratão-lavadeiro (‘washing rat’) in French and Portuguese, respectively. So, whether the raccoon is related to bears, dogs, or rats, every European taxonomist of the 18th century agreed that washing was its defining trait. But it is a popular misconception.

Added in the revised version: Raccoons wet their foods to enhance haptic sensation, not eliminate contaminants;their intent is to douse food objects, not wash them (Lyall-Watson, 1963). Intriguingly, the posterior cerebrum—which includes the somatosensory cortex and several other cortical, thalamic, and subcortical structures—is relatively expanded among raccoons, suggesting magnification of the neural pathways that serve tactile processing (Arsznov and Sakai, 2013). When viewed in this light, the Lenape ethnonym nachenum (‘they use hands as tools’) would seem most fitting of all, not least for blurring the thin line between human and nonhuman curiosity and the nature of knowing.

Earlier version: We conducted 101 feeding trials, recording 1,282 food-handling events by 42 individuals. We found that monkeys were sensitive to sand on their food, responding to each treatment—low, intermediate, and high concentrations—with greater median durations (± 1 SD) of brushing (low: 0.0 ± 0.1 s; intermediate: 1.1 ± 2.0 s; high: 3.1 ± 2.0 s; Figure 2B) and washing (low: 0.04 ± 0.3 s; intermediate: 0.6 ± 2.0 s; high: 3.3 ± 4.3 s; Figure 2C). This result is important for upholding long-held assumptions of intentional cleaning. Further, we found that dominant monkeys of both sexes showed a strong propensity for food-brushing (Figure 2B; Table S3) over food-washing (Figure 2C; Table S4). This finding reverses the pattern observed on Koshima Island (Watanabe, 1994), and it raises the possibility that food-washing is an indulgence subject to diminishing returns. To explore this premise, we developed a theoretical model where the time devoted to food-cleaning is predicted to maximise the rate of sand removal as a function of handling time.

Revised version: We conducted 101 feeding trials, recording 1,282 food-handling events by 42 individuals. We had detailed rank information for 23 individuals, so we used this subset of data in our GLMM analyses. We found that all monkeys were sensitive to sand on their food, responding to each treatment—low, intermediate, and high concentrations—with greater median durations (± 1 SD) of brushing (low: 0.0 ± 0.1 s; intermediate: 1.1 ± 2.0 s; high: 3.1 ± 2.0 s; Figure 2B) and washing (low: 0.04 ± 0.3 s; intermediate: 0.6 ± 2.0 s; high: 3.3 ± 4.3 s; Figure 2C). This result is important for upholding long-held assumptions of intentional cleaning. Further, we found that dominant monkeys of both sexes showed a strong propensity for food-brushing (Figure S2A; Table S3) over food-washing (Figure S2B; Table S4). This finding reverses the pattern observed on Koshima Island (Watanabe, 1994), and it raises the possibility that food-washing is an indulgence subject to diminishing returns. To explore this premise, we developed a theoretical model where the time devoted to food-cleaning is predicted to maximise the rate of sand removal as a function of handling time.

Earlier version: Figure 3A illustrates the fastidious nature of our study population: monkeys allocated excess time to washing and brushing—by factors of 1.5 and 3.0, respectively—beyond that predicted by the optimization of sand removal. Our model also highlights sharply divergent responses to the sunk costs of food-handling time (Figure 3B). Given the efficacy of washing (Figure S1) and time needed to carry food to the ocean ( x¯ = 22 ± 15 s; range: 5 to 78 s), there is little incentive to over-wash food (Figure 3B, region I). At the same time, the lowest- and highest-ranking monkeys abstained from washing altogether, choosing instead to minimize food-handling time by overbrushing their food (Figure 3B, region II). This tolerance for fast-diminishing returns underscores the monkeys’ strong aversion to sand; but even so, the long-term benefits of mitigating tooth wear must be balanced against urgent energetic requirements.

Revised version: Figure 3A illustrates the fastidious nature of our study population: monkeys allocated excess time to washing and brushing—by factors of 1.5 and 3.0, respectively—beyond that predicted by the optimization of sand removal Video 2). Our model also highlights sharply divergent responses to the sunk costs of food-handling time (Figure 3B). Given the greater efficacy of washing (Figure S1) and time needed to carry food to the ocean ( x¯ = 22 ± 15 s; range: 5 to 78 s), there is little incentive to over-wash food (Figure 3B, region I). At the same time, the lowest- and highest-ranking monkeys abstained from washing altogether, choosing instead to minimize food-handling time by over-brushing their food (Figure 3B, region II). This tolerance for fast-diminishing returns underscores the monkeys’ strong aversion to sand; but even so, the long-term benefits of mitigating tooth wear must be balanced against urgent energetic requirements.

Predicted and observed cleaning times.

Earlier version: a, Predicted time (large filled points) vs. observed times (violin plots) for brushing and washing food (note log scale). b, Predicted cleaning time as a function of cleaning inefficiency c, and handling time h, with predicted values (black points) for brushing and washing based on observed cleaning inefficiencies and handling times. The coloured points (as in panel a) represent observed cleaning times. The tradeoff between longer food handling times and efficient cleaning (oceanside food-washing; Region I) and shorter handling times and inefficient cleaning (immediate food-brushing; Region II) is depicted by the black curve.

Revised version: (A) Mean predicted time (large filled points vs. observed times (violin plots) for brushing and washing food (note log scale). The vertical line associated with predicted times represents the 5 to 95% confidence interval. (B) Predicted cleaning time as a function of cleaning inefficiency c, and handling time h, with mean predicted values (black points) for brushing and washing based on observed cleaning inefficiencies and handling times. The colored points (as in panel A) represent observed cleaning times. The trade-off between longer food handling times and efficient cleaning (oceanside food-washing; Region I) and shorter handling times and inefficient cleaning (immediate food-brushing; Region II) is depicted by the black curve.

Disposable soma

The disposable-soma hypothesis of senescence predicts investment in the immediate needs of survival or reproduction over tooth preservation (Carranza et al., 2004). Dominant monkeys face this predicament because rapid food intake rates are integral to sustaining dominance and accruing reproductive success. For dominant males, energy intake determines their ability to sustain consortships during mating (Higham et al., 2011); and, for dominant females, it affects practically every measure of fitness (Alberts, 2019; Cooper et al., 2022). Our findings suggest that dominant monkeys refrained from washing to maximize short-term energy intake (Figure 2D). In short, they prioritized pressing energetic needs over the long-term benefits of tooth preservation—a ‘live fast, die young’ life-history strategy. This view of teeth as disposable soma may explain why dominant male monkeys experience faster senescence and earlier mortality (Anderson et al., 2021). Estimating fitness consequences is beyond the scope of our study, but our findings suggest that a prolonged life is also subject to diminishing returns.

Paleo matter(s)

The full extent of sand-mediated tooth wear is unknown for our study population, but it is probably extreme among the highest-ranked individuals. If affirmed, the findings could affect our views of the hominin fossil record by challenging the assumption that dietary variability is the principal cause of variable dental wear. Some species, notably Paranthropus boisei, had ready access to water, which raises the possibility that they—like many primate species—assiduously washed their food, an essential behavior if their diet featured gritty underground plant tissues (Wrangham et al., 2009; Fannin et al., 2021). Other species, notably P. robustus, have extremely variable levels of tooth pitting (Peterson et al., 2018), which could reflect, at least partially, the absence of extensive wetlands (Herries et al., 2010) coupled with interindividual variation in food-cleaning behaviors. Tellingly, the dental wear observed on the macaques of Koshima Island bears striking similarities to the hominin fossil record (Towle et al., 2022), suggesting that populations of food-cleaning monkeys are a valuable model system that warrant further study.

Significance

Earlier version: Our study leverages a new method in ecological research to provide the first analysis of siliceous particles on primate foods. Our experiment probes the behavioral economics of wild monkeys, revealing a strong aversion to sandy foods. Yet, the monkeys behaved irrationally when cleaning their foods, allocating excess time than predicted by an optimization model. Some individuals fell victim to the sunk cost fallacy by over-washing their foods, whereas dominant monkeys abstained from washing altogether, seemingly sacrificing their teeth at the altar of high rank, a social status that relies rapid food intake. Our results support the disposable-soma hypothesis for senescence and kick the tires of a treasured assumption in paleoanthropology.

Revised version: Our study leverages a new method in ecological research to provide the first analysis of siliceous particles on primate foods. Our experiment probes the behavioral economics of wild monkeys, revealing a strong aversion to sandy foods. Yet, the monkeys behaved irrationally when cleaning their foods, allocating excess time than predicted by an optimization model. Some individuals fell victim to the sunk cost fallacy by over-washing their foods (Box 2), whereas dominant monkeys abstained from washing altogether, seemingly sacrificing their teeth at the altar of high rank, a social status that relies rapid food intake. Our results support the disposable-soma hypothesis for senescence while kicking the tires of a treasured assumption in paleoanthropology.

Added in the revised version: Box 2. Balancing sunk costs against future benefits

Added in the revised version: The Aérospatiale-BAC Concorde 102 is an iconic aircraft manufactured from 1965-1979. This photo shows British Airways flight 002 on the eve of its final commercial flight on October 24, 2003 (photograph by Richard Vandervord, reproduced with permission).

Added in the revised version: Carrying food to the ocean costs time and energy. Some monkeys stood upright and walked to the water because their hands held cucumbers, paying an extremely high energetic cost (Nakatsukasa et al., 2006). Such irrecoverable expenses—or ‘sunk costs’—should not sway the optimal washing time of rational monkeys, which we calculated as 2.40 ± 0.74 sec per cucumber slice. Still, many monkeys washed their food far beyond the point of diminishing returns (Figure 3A). This level of over-washing speaks to their aversion to sand, but it also suggests that some individuals succumbed to the sunk cost fallacy, or “Concorde Effect” (Arkes and Ayton, 1999).

Added in the revised version: The Concorde is a supersonic aircraft. Production began in 1962 as a joint Anglo-French enterprise, but significant cost overruns undermined any chance of profit. By 1971, the opening line of a British central policy review staff memorandum was explicit and brutal: “Concorde is a commercial disaster,” it said. “It should never have been started.” Still, both governments continued to pour millions into the project on the grounds that they had already placed substantial investment in it. Such reasoning is viewed as an economic fallacy because decisions should be based on prospective (future) costs only. The same memo put it this way, “The decision whether or not to abandon Concorde must start from where we are now—much of the milk is already spilt.”

Added in the revised version: Irrational decisions are not unique to humans. When faced with a foraging task that required inaction, humans, rats, and mice each fell victim to the Concorde Effect; they were more likely to complete a trial (i.e., to continue waiting instead of opting out) the longer they had already waited (Sweis et al., 2018). In another experiment, Watzek and Brosnan (2020) showed that tufted capuchins and rhesus macaques are reluctant to forfeit a small investment when completing a psychomotor task, persisting 5–7 times longer than was optimal for a food reward. A criticism of these findings is that animal hunger can bias the results, impelling irrational decisions (Ott et al., 2022). In our experiment, however, middle-ranking monkeys delayed consumption of a food reward already in their possession, suggesting an incentive to offset the long-terms costs of sand-mediated tooth damage. If irrational over-washing confers adaptive benefits by prolonging the functional life of teeth, then there could be a kernel wisdom in the Concorde Effect.

Methods and materials

Study site and population

Earlier version: Koram Island (12.242°, 100.009°) lies ~1 km offshore in the Gulf of Thailand and within Khao Sam Roi Yot National Park, Prachuap Khiri Khan, Thailand. It has an area of 0.45 km2 and a coastline of 3.5 km. The habitat—limestone karst blanketed with a dense flora of dwarf evergreen trees and deciduous scrub, and encircled by rocky shore and sandy beaches—supports a population of ca. 75 long-tailed macaques described as hybrids at the subspecies taxonomic level (Macaca fascicularis aurea x M.f. fascicularis)(Gumert et al.,2019). The animals are well habituated to human observers due to regular tourism and sustained study since 2013 (Tan et al., 2018).

Revised version: Koram Island (12.242°, 100.009°) lies ~1 km offshore in the Gulf of Thailand and within Khao Sam Roi Yot National Park, Prachuap Khiri Khan, Thailand. It has an area of 0.45 km2 and a coastline of 3.5 km. The habitat—limestone karst blanketed with a dense flora of dwarf evergreen trees and deciduous scrub, and encircled by rocky shore and sandy beaches—supports a population of ca. 75 long-tailed macaques described as hybrids at the subspecies taxonomic level (Macaca fascicularis aurea × M.f. fascicularis)(Gumert et al., 2019). The animals are well habituated to human observers due to regular tourism and sustained study since 2013 (Tan et al., 2018). Most of this research has revolved around stone tool-mediated foraging on mollusks, the only activity known to elicit stone handling (Malaivijitnond et al., 2007; Gumert and Malaivijitnond, 2012, 2013; Tan et al., 2015), although infants and juveniles will sometimes use stones during object play (Tan, 2017). There has been no prior examination of food-cleaning behaviors.

Rank determination

Earlier version: Macaques form multi-male multi-female (polygynandrous) social groups with individual dominance hierarchies. Among females, this hierarchy is strictly linear and stable through time (van Noordwijk and van Schaik, 1999). To determine the rank-order of adults, we recorded dyadic agonistic interactions and their outcomes (i.e., aggression, supplants, and silent-bared-teeth displays of submission) during 5-min focal follows of individuals based on a randomized order of continuous rotation (Tan et al., 2018). In some cases, these data were supplemented with ad libitum observations. This protocol existed during five years (2013-2018) of continual observations before we conducted our experiment in July-August 2018. To determine the effects of dominance rank on individual foodcleaning propensities, and to standardize the data by sex, we followed the methods of Levy et al. (2020). We calculated ordinal ranks between 1 (highest) and n (lowest), where n is the number of animals aged ≥ 5 years in each hierarchy. We analyzed male (n = 8) and female (n = 16) dominance hierarchies together in the same statistical models. We chose ordinal ranks because they best approximate competitive regimes during density-dependent competition (Levy et al., 2020). This type of feeding competition reflects those of our experimental trials and the natural conditions on Koram Island, where preferred resources are limited (Luncz et al., 2017).

Revised version: Macaques form multi-male multi-female (polygynandrous) social groups with individual dominance hierarchies. Among females, this hierarchy is strictly linear and stable through time (van Noordwijk and van Schaik, 1999). To determine the rank-order of adults, we recorded dyadic agonistic interactions and their outcomes (i.e., aggression, supplants, and silent-bared-teeth displays of submission) during 5-min focal follows of individuals based on a randomized order of continuous rotation (Tan et al., 2018). In some cases, these data were supplemented with ad libitum observations. This protocol existed during five years (2013–2018) of continual observations before we conducted our experiment in July-August 2018. To determine the effects of dominance rank on individual food-cleaning propensities, we followed the methods of Levy et al. (2020). We calculated ordinal ranks between 1 (highest) and n (lowest), where n is the number of animals aged ≥ 5 years. Then we combined males (n = 8) and females (n = 16) into a single standardized ordinal ranking of 24 animals. This method of rank determination is well suited for conditions involving densitydependent competition (Levy et al., 2020), such as those of our experiment and more generally on Koram Island, where preferred food resources are limited (Luncz et al., 2017).

Measuring sand

To quantify the amount of sand on food surfaces, whether provisioned by tourists (cucumbers, melon, pineapple) or used in our experiment (sliced cucumbers), we applied a quick-drying liquid polymer—granulated plastic (Pioloform BL 16; Wacker-Chemie GMBH, Munich, Germany) mixed with ethanol (18% plastic to 82% ethanol, by weight)—to each food item. When dried, we peeled and stored the sand-infused film for analysis. The advantages of this method are twofold: the removal of exogenous particulate matter is extremely thorough; and the plastic does not detach biogenic silica such as trichromes or phytoliths (Hinton et al., 1996). In the lab, we dissolved each peel in ethanol and separated the sand by centrifugation, producing a pellet. We dried and weighed the pellet, dividing the mass by the surface area of the food object, which we calculated from digital photographs imported into ImageJ v. 1.52. This method produces an estimate of exogenous particulate mass per area (mg mm-2), allowing direct comparison of apples and oranges.

Earlier version: To measure the elemental and physical properties of sand, we dispersed and filtered the pellets in water using a 0.2-μm isopore membrane filter, which we submitted for scanning electron microscopy (SEM) and energy dispersive X-ray spectroscopy (EDS). To establish the parameters for multi-field and bulk analysis, we imaged a representative area of the filter at multiple magnifications and performed discrete particle analysis. 50× magnification allowed for statistically significant representation of particle number and size range (allowing a 5 μm lower particle size range in analysis). All discrete particle analyses indicated silicon rich particles and composition distribution bins were established to include dominant accompanying elements. After establishing these parameters, we initiated multi-field automated analysis using six fields of view of the debris field (at 50×). A composition classification was assigned to each particle and data sorted by composition classification and particle size (particle sizing was binned using standard Feret maximum parameter). The sizing bins are standard ISO-16232 size classes.

Revised version: To measure the elemental and physical properties of sand, we dispersed and filtered the pellets in water using a 0.2-μm isopore membrane filter, which we submitted for scanning electron microscopy (SEM) and energy dispersive X-ray spectroscopy (EDS). To establish the parameters for multi-field and bulk analysis, we imaged a representative area of the filter at multiple magnifications and performed discrete particle analysis. 50× magnification allowed for statistically significant representation of particle number and size range (allowing a 5 μm lower particle size range in analysis). All discrete particle analyses indicated silicon rich particles and composition distribution bins were established to include dominant accompanying elements. After establishing these parameters, we initiated multi-field automated analysis using six fields ofviewof the debris field (at 50×). A composition classification was assigned to each particle and data sorted by composition classification and particle size (particle sizing was binned using standard Feret maximum parameter). The sizing bins are standard ISO-16232 size classes.

Experimental design

Earlier version: To elicit food handling behaviors and determine individual cleaning preferences, we put three cucumber slices in each of three trays (20×30×10 cm) and manipulated the amount of contaminating sand. In the low-sand treatment, we put cucumber slices in a tray without sand; however, the presence of some aeolian sand was unavoidable. In the intermediate-sand treatment, we lined the tray with sand and put the cucumber slices on the surface. In the high-sand treatment, we buried cucumber slices completely (Figure 2A). Trays were placed 1.5 m apart and 15 m from the ocean, and we randomized the color and sequence of trays across trials.

Revised version: To elicit food handling behaviors and determine individual cleaning preferences, we put three cucumber slices in each of three trays (20 ×30 × 10 cm) and manipulated the amount of contaminating sand. In the low-sand treatment, we put cucumber slices in a tray without sand; however, contamination from aeolian sand was unavoidable. In the intermediate-sand treatment, we lined the tray with sand and put the cucumber slices on the surface. In the high-sand treatment, we buried cucumber slices completely (Figure 2A). Trays were placed 1.5 m apart and 15 m from the ocean, and we randomized the color and sequence of trays across trials.

Earlier version: Trials began when one or more monkeys approached the trays and ended when the animals finished every cucumber slice or abandoned the experiment (range: 10 s to 14 min). We used video recordings to determine the onset and offset of individual food-handling bouts, beginning from initial contact with a cucumber slice and ending when the final slice, in its entirety, entered the mouth. Within each bout, we determined the duration of brushing and washing behaviors, defining each from the onset of serial stereotypical forelimb movements to the moment of oral ingestion. We estimated energy intake rates by calculating the number of cucumber slices consumed during each food-handling bout, and multiplying each slice by 1.1 kcal (source: U.S. Department of Agriculture, FoodData Central, 2019) and 4.186 kJ (Hargrove, 2007). We performed 101 trials over 5 weeks and recorded 1,282 food-handling bouts by 42 individual monkeys. To minimize the potential confounding effects of dominance interactions, we analyzed trials with ≤ 3 monkeys. Thus, 935 food-handling bouts were included in GLMM statistical models, which included data on individual rank, sex, and sand treatment. Ifa monkey consumed a cucumber slice without brushing or washing it, the zero-second duration was included in both GLMMs.

Revised version: Trials began when one or more monkeys approached the trays and ended when the animals finished every cucumber slice or abandoned the experiment (range: 10 s to 14 min). We used video recordings to determine the onset and offset of individual food-handling bouts, beginning from initial contact with a cucumber slice and ending when the final slice, in its entirety, entered the mouth. Within each bout, we determined the duration of brushing and washing behaviors, defining each from the onset of serial stereotypical forelimb movements to the moment of oral ingestion. We estimated energy intake rates by calculating the number of cucumber slices consumed during each food-handling bout, and multiplying each slice by 1.1 kcal (source: U.S. Department of Agriculture, FoodData Central, 2019) and 4.186 kJ (Hargrove, 2007). We performed 101 trials over 5 weeks and recorded 1,282 food-handling bouts by 42 individual monkeys. To minimize the potential confounding effects of dominance interactions, we analyzed trials with ≤ 3 monkeys. Thus, 937 food-handling bouts were included in GLMM statistical models, which included data on individual rank, sex, and sand treatment. Ifa monkey consumed a cucumber slice without brushing or washing it, the zero-second duration was included in both GLMMs.

Behavioral analyses

Earlier version: To model experimental variance in brushing and washing behaviors as a function of experimental treatment, sex, and rank, we fit generalized linear mixed models (GLMMs) using the glmmTMB package in R version 4.2.3 (Brooks et al., 2017). In each GLMM, we modelled sand treatment (a categorical variable with three levels), sex (a categorical variable with two levels), and ordinal rank (a discrete variable ranging from 1 to 18) as fixed effects. We also incorporated two additional interaction terms as fixed effects: sand treatment χ ordinal rank and sand treatment χ sex. To account for experimental variance among individuals and control for pseudoreplication (because the number of feeding bouts per individual varied widely; Bolker et al. (2009)), we included individual ID as a random intercept. The brushing and washing data sets were whole-number counts (seconds) with means < 5. The distributions were right-skewed with high concentrations of biologically-meaningful zeros (Martin et al., 2005) (i.e., instances of food-handling without any cleaning behavior). Thus, we fit four separate models in glmmTMB to account for these non-normal spreads: standard GLMMs with a log-link function and either a (1) Poisson or (2) negative binomial error distribution (default = nbinom2 in glmmTMB); or zero-inflated generalised linear mixed models (ZIGLMM) with a logit-link function, a single zero-inflation parameter applying to all observations, and either a (3) Poisson or (4) negative binomial error distribution. We then determined the best fit model using delta AIC values in the bbmle package in R.

Revised version: To model variance in brushing and washing behaviors as a function of experimental treatment and rank, we fit generalized linear mixed models (GLMMs) using the glmmTMB package in R version 4.4.1 (Brooks et al., 2017). We then evaluated the performance of each model using the standardized residual diagnostic tools for hierarchical regression models available in the DHARMa package (Hartig, 2022). In each GLMM, we modelled sand treatment (a categorical variable with three levels), sex (a categorical variable with two levels), and ordinal rank (a discrete variable ranging from 1 to 24) as fixed effects. For our brushing model, we incorporated an additional interaction term as a fixed effect: sand treatment × ordinal rank. For our washing dataset, we ran two models: the first modeled the relationship between ordinal rank and food cleaning as a linear term, while the second modeled the relationship as a quadratic term. We also incorporated both the linear and quadratic terms into the interaction effect with sand treatment in each washing model. To account for experimental variance among individuals and control for pseudoreplication (because the number of feeding bouts per individual varied widely; Bolker et al. (2009)), we included individual ID as a random intercept.

Added in the revised version: Our brushing and washing data sets were whole-number counts (seconds) with means < 5. The distributions were right-skewed with high concentrations of biologically-meaningful zeros (Martin et al., 2005) (i.e., instances of food-handling without any cleaning behavior). Thus, we fit a series of zero-inflated generalized linear mixed models (ZIGLMM), each with a logit-link function, a single zero-inflation parameter applying to all observations, and Poisson error distribution. For the foodbrushing model, we fit a zero-inflated Poisson (ZIP), which produced favorable standardized residual diagnostic plots with no major patterns of deviation (Figure S3) and minor, but non-significant, underdispersion (DHARMa dispersion statistic = 0.99, p = 0.80). For our two food-washing models, we used zero-inflated models with Conway-Maxwell Poisson (ZICMP) distributions, an error distribution chosen for its ability to handle data that is more underdispersed (DHARMa dispersion statistic = 8.2E-09, p = 0.74) than the standard zero-inflated Poisson (Brooks et al., 2019). Using this error distribution improved residual diagnostic plots over a standard ZIP model and we view any deviations in the standardized residuals as minor and due to the smaller sample size of our food-washing data set (Figures S4 and S5) (Hartig, 2022). We report the summarized fixed effects tests for each GLMM in Tables S1–S3 as Analysis of Deviance Tables (Type II Wald chi square tests, one-sided) along with χ2 values, degrees of freedom, and p-values (one-sided tests). Full model summaries with standard errors and confidence intervals are also included in Tables S4–S6. For all statistical analyses, we set α = 0.05.

Removed in the revised version: dAIC values indicated that, for both brushing (Table S1) and washing (Table S2), negative binomial GLMMs without zero-inflation were the best-fitting models. We validated each model by calculating dispersion statistics (χ2 /degrees of freedom) (Zuur and leno, 2016). Dispersion statistics for the brushing model (χ2 = 541.9; residual degrees of freedom = 564; χ2 /rdf = 0.96, p = 0.74, one-sided test) and the washing model (χ2 = 174.2; residual degrees of freedom = 351; χ2 /rdf = 0.50, p = 1.00, one-sided test) failed to detect overdispersion in either case. We report the fixed effects tests for each GLMM in Tables S3 and S4 as Analysis of Deviance Tables (Type II Wald chi square tests, one-sided) along with χ2 values, degrees of freedom, and p-values (one-sided tests). For all statistical analyses, α = 0.05.

Optimal cleaning time model

Earlier version: To model the optimization of sand removal before consumption, we accounted for two distinct temporal periods: handling time h, which includes an assessment time and pre-cleaning time, and the cleaning time t. Assessment time (set as a constant 1 second) includes visual fixation on a food object and forelimb extension before contact, whereas pre-cleaning time represents all handling activities that precede cleaning. During brushing, the pre-cleaning time was essentially nil (zero seconds), but washing required travel from the experimental treatments to the ocean, requiring longer pre-cleaning times ( x¯ = 22 ± 15 s; range: 5 to 78 s). We assumed that the proportion of sand removed from each cucumber follows the saturating relationship g(t) = t/(c +1), where c is the halfsaturation constant associated with brushing or washing. As c increases, so does the inefficiency of a given cleaning behaviour. Given our observations that brushing removes 75% of grit in 2.97 s, and washing removes 93% of grit in 3.53 s (Figure S1), we obtain the constants cbrushing = 0.99 s and cwashing=0.26 s, such that washing (without considering handling costs) is the most efficient strategy. The rate of grit removal is then given by R(t) = g(t)/(h + t), which reaches a maximum at the optimal cleaning time t*=ch . For brushing and washing cleaning strategies, we obtain the expected optimal cleaning times t*brushing = 0.98 s, and t*washing = 2.39 s (Figure 3a), respectively. These optimal cleaning times are defined exclusively with respect to maximising the rate of grit removal, without considering the potentially cascading effects of these strategies on fitness.

Revised version: To model the optimization of sand removal, we drew inspiration from the marginal value theorem of Charnov (1976), defining two temporal periods: handling time h, which includes an assessment time and pre-cleaning time, and cleaning time t. Assessment time (set as a constant 1 second) includes visual fixation on a food object and forelimb extension before contact, whereas pre-cleaning time represents all handling activities that precede cleaning. During brushing, the pre-cleaning time was essentially nil (zero seconds), but washing required travel from the experimental treatments to the ocean, requiring longer pre-cleaning times ( x¯ = 22 ± 15 s; range: 5 to 78 s). We assumed that the proportion of sand removed from each cucumber follows the saturating relationship g(t) = t/(c +t), where c is the cleaning inefficiency, or the half-saturation constant associated with brushing or washing. As c increases, so does the inefficiency of a given cleaning behavior. Given our observations, it requires on average 2.97 s of brushing to remove 75±7% of grit, and 3.53 s of washing to remove 93 ± 4% of grit (Figure S1). From these data, and including the experimental uncertainty associated with grit removal percentages, we obtained distributions for estimated cleaning inefficiencies, cbrushing = 0.99 ± 0.38 s and cwashing = 0.27 ± 0.15 s, such that washing (without considering handling costs) is the most efficient strategy. The rate of grit removal is then given by R(t) = g(t)/(h +t), which reaches a maximum at the optimal cleaning time t*=ch . For brushing and washing cleaning strategies, we obtain the expected optimal cleaning times tbrushing* = 0.98±0.19 s, and twashing* = 2.40 ± 0.74 s (Figure 3a), respectively, and where including additional sources of uncertainty did not alter our findings (Figure S6). These optimal cleaning times are defined exclusively with respect to maximizing the rate of grit removal, without considering the potentially cascading effects of these strategies on fitness.